**Q1**: Implement the function named `even_or_odd(num)` that checks if the number `num` is odd or even, and returns the string "odd" or "even" accordingly.

<!-- grade_id: cell-70432f0a5b317430, points: 0, grade: false, solution: true -->
```python
def even_or_odd(num):
  # YOUR CODE HERE
  return num % 2 == 0 and "even" or "odd"
```

<!-- grade_id: cell-01ff4c49b163bc43, points: 1, grade: true, solution: false -->
```python
# Test cases
assert even_or_odd(0) == "even"
assert even_or_odd(2) == "even"
assert even_or_odd(-2) == "even"
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-82342b52b9d30d79, points: 1, grade: true, solution: false -->
```python
# Test cases
assert even_or_odd(1) == "odd"
assert even_or_odd(-1) == "odd"
print('ok')
```

Output:
```
ok
```

**Q2**: Implement the function named `what_animal(characteristics)` that identifies the animal class from the given characteristics string using the phylogenetic tree below, and returns the animal class.

![animal-phylogenetic-tree.png](https://drive.google.com/uc?id=1rbVbNugM3NKaI4sbsPF5uMatyE-ZpdKc)

Example:

-----Argument-----\
'Vertebral column, Hinged jaw, Legs, Egg with ammnion'\
-----Return--------\
'Lizard'

-----Argument-----\
'Vertebral column, Legs, Hinged jaw'\
-----Return--------\
'Frog'


<!-- grade_id: cell-7ed364d819021dc6, points: 0, grade: false, solution: true -->
```python
def what_animal(characteristics):
  # YOUR CODE HERE
  characteristics = characteristics.split(", ")

  if 'Vertebral column' in characteristics:
    if 'Hinged jaw' in characteristics:
      if 'Legs' in characteristics:
        if 'Egg with ammnion' in characteristics:
          if 'Hair' in characteristics:
            return 'Rabbit'
          else: return 'Lizard'
        else: return 'Frog'
      else: return 'Fish'
    else: return 'Lamprey'
  else: return 'Lancelet'
```

<!-- grade_id: cell-c89af4a9263520b3, points: 1, grade: true, solution: false -->
```python
# Test cases
assert(what_animal('Vertebral column, Hinged jaw, Legs, Egg with ammnion, Hair')== 'Rabbit')
assert(what_animal('Hinged jaw, Vertebral column, Legs, Egg with ammnion, Hair')== 'Rabbit')
assert(what_animal('Vertebral column, Egg with ammnion, Hinged jaw, Legs')== 'Lizard')
assert(what_animal('Vertebral column, Hinged jaw, Legs, Egg with ammnion')== 'Lizard')
assert(what_animal('Vertebral column, Hinged jaw, Legs')== 'Frog')
assert(what_animal('Vertebral column, Hinged jaw')== 'Fish')
assert(what_animal('Vertebral column')== 'Lamprey')
assert(what_animal('')== 'Lancelet')
print('ok')
```

Output:
```
ok
```

**Q3**: Implement the function named `factorial(n)` that returns the factorial of a number `n` using `while` or `for`.

<!-- grade_id: cell-aaf79f77e9dfdd37, points: 0, grade: false, solution: true -->
```python
def factorial(n):
  # YOUR CODE HERE
  if n == 0:
    return 1
  else:
    return n * factorial(n-1)
```

<!-- grade_id: cell-ed52121c72bd381b, points: 1, grade: true, solution: false -->
```python
# Test cases
assert factorial(0) == 1
assert factorial(1) == 1
assert factorial(2) == 2
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-8bfeb4d19a1f83cb, points: 1, grade: true, solution: false -->
```python
# Test cases
assert factorial(3) == 6
assert factorial(4) == 24
assert factorial(5) == 120
print('ok')
```

Output:
```
ok
```

**Q4** : Implement the function named `search_sub_seq(dna_seq, sub_seq)` that searches for a nucleotide subsequence `sub_seq` in a given DNA sequence `dna_seq`, and returns the position of the first nucleotide subsequence if found or `None` otherwise. Note that the position of the nucleotide sequence starts at 1

Example: 

-----Argument-----\
"GGTACGTACGTACGTACGT", "ACGT"\
-----Return--------\
4

<!-- grade_id: cell-fad84e9a9c55d63a, points: 0, grade: false, solution: true -->
```python
def search_sub_seq(dna_seq, sub_seq):
  # YOUR CODE HERE
  for i in range(len(dna_seq)):
    if dna_seq[i:i+len(sub_seq)] == sub_seq:
      return i+1
  return None
```

<!-- grade_id: cell-3d173c04086a2a94, points: 1, grade: true, solution: false -->
```python
# Test cases
assert(search_sub_seq("ATGC", "TTT")== None)
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-322745d8524ab45c, points: 1, grade: true, solution: false -->
```python
# Test cases
assert(search_sub_seq("ATGC", "C")== 4)
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-4e8341c8eb8be23c, points: 1, grade: true, solution: false -->
```python
# Test cases
assert(search_sub_seq("ATGC", "GC")== 3)
assert(search_sub_seq("GGTACGTACGTACGTACGT", "ACGT")== 4)
print('ok')
```

Output:
```
ok
```

**Q5**: Implement the function named `find_max_hypotenuse(perimeter)` that finds the maximum integer length of the hypotenuse of any right triangle subject to the perimeter length received. Note that we consider only those right triangles whose side lengths are all integers.
For instance, if the given perimeter is 90, there are two possible triangles that meet the requirements: 15, 36, 39 and 9, 40, 41. Thus, the valure returned from this function is 41.

![pythagoras.png](https://drive.google.com/uc?id=1vjYjrD8T35vkg_scMInIUZDn8J8bhaiJ)

Example: 

-----Argument-----\
90\
-----Return--------\
41\
-----Argument-----\
30\
-----Return--------\
13

<!-- grade_id: cell-36e24ba46b84fad3, points: 0, grade: false, solution: true -->
```python
def find_max_hypotenuse(perimeter):
  # YOUR CODE HERE

  for i in range(1, perimeter//2):
    for j in range(1, perimeter//2):
      k = perimeter - i - j
      if i**2 + j**2 == k**2:
        return k
  return None

  
```

<!-- grade_id: cell-48cc53ab925270a2, points: 1, grade: true, solution: false -->
```python
# Test cases
assert(find_max_hypotenuse(5)== None)
assert(find_max_hypotenuse(1)== None)
print('ok')
```

Output:
```
ok
```

<!-- grade_id: cell-5c12870de2a0a936, points: 1, grade: true, solution: false -->
```python
# Test cases
assert(find_max_hypotenuse(12)== 5)
assert(find_max_hypotenuse(90)== 41)
assert(find_max_hypotenuse(30)== 13)
print('ok')
```

Output:
```
ok
```