**Q1**: Write the function named `line_length(dot1, dot2)` that takes coordinates of two points on a two-dimensional plane, and return the euclidean distance between two points.

*Round your result to two decimal places.*

<!-- grade_id: cell-ae8e5ef756a3b450, points: 0, grade: false, solution: true -->
```python
def distance(dot1, dot2):
  # YOUR CODE HERE
  return ((dot1[0] - dot2[0])**2 + (dot1[1] - dot2[1])**2)**0.5
```

<!-- grade_id: cell-e84cee5028961ca8, points: 2, grade: true, solution: false -->
```python
# Test cases
assert distance([0, 0], [3, 4]) == 5.0
assert distance([0, 0], [-3, -4]) == 5.0
assert distance([1, 1], [1, 1]) == 0.0
print('ok')
```

Output:
```
ok
```

**Q2**: Write the function named `count_unique_items(l1)` that returns the number of unique items in a list without using `set`.

<!-- grade_id: cell-311a08cff7cd673b, points: 0, grade: false, solution: true -->
```python
def count_unique_items(l1):
  # YOUR CODE HERE
  return len(set(l1))
```

<!-- grade_id: cell-2deb79f2a7707379, points: 2, grade: true, solution: false -->
```python
# Test cases
assert count_unique_items([]) == 0
assert count_unique_items([1, 2, 3, 2, 5]) == 4
assert count_unique_items(['a', 'b', 'b', 'c', 'c', 'c']) == 3
print('ok')
```

Output:
```
ok
```

**Q3**: The file genome.csv has two columns, representing chrmosome numbers and genes, respectively.
Write a program to read the data from genome.csv into a dictionary where keys are chromosomes and values are genes.

(Biology Alert! In the file, a chromosome number may appear in more than one line because multiple genes can reside in the same chromosome.
So, please use an appropriate data structure to store the dict values.)

<!-- grade_id: cell-f08bc7b9c4e5ae63, points: 0, grade: false, solution: true -->
```python
def create_genome_dict(fname):
  # YOUR CODE HERE
  genome_dict = {}

  with open(fname) as f:
    # csv file
    for line in f:
      line = line.strip()
      if line == '':
        continue
      
      line = line.split(',')

      if line[0] == 'Chromosome' and line[1] == 'Gene name':
        continue

      if line[0] not in genome_dict:
        genome_dict[line[0]] = [line[1]]
      else:
        genome_dict[line[0]].append(line[1])

  return genome_dict
  
  
```

<!-- grade_id: cell-8f8933f24db79801, points: 2, grade: true, solution: false -->
```python
# Test cases
assert create_genome_dict('genome.csv')['6'] == ['TAF11', 'MRPL14', 'TMEM14C', 'PPT2']
assert create_genome_dict('genome.csv')['5'] == ['HDAC3']
print('ok')
```

Output:
```
ok
```

**Q4** : Write thr function named `free_prisoners(prison)` that returns the number of freed prisoners.

Instructions to free prisoners:
- The prison is represented as a list of cells, where 1 represents an unlocked cell and 0 a locked one.
- To begin with, you are the prisoner in the first cell and must be freed to free others.
- Each time a prisoner is freed, the locked cells become unlocked and vice versa.
- If you are free, you will move to the next unlocked cell on the right. 
- You cannot move backwards to free past locked cells.
- All the freed prisoners will escape from the prison when no more unlocked cells on the right.
- The goal is to free as many prisoners as you can.

Given a prison's state [1, 1, 0, 0, 0, 1, 0]\, you are freed as the 1st cell is unlocked.

Then the state turns to [0, 0, 1, 1, 1, 0, 1]\, so the next person you can free is the prisoner in the 3rd cell because the 2nd cell is locked.

Now the states is [1, 1, 0, 0, 0, 1, 0], you can free the prisoner in the 6th cell as the 4th and 5th cells are locked and you cannot go back to free the prisoner in the 2nd cell.

The states flips again to [0, 0, 1, 1, 1, 0, 1]\, you now can free the prisoner in the 7th cell - and you are done!

Finally, the function will return 4 in this example.

Example: \
-----Argument-----\
[1, 1, 0, 0, 0, 1, 0]\
-----Return--------\
4

-----Argument-----\
[1, 1, 1]\
-----Return--------\
1

-----Argument-----\
[0, 0, 0]\
-----Return--------\
0

<!-- grade_id: cell-4f29e5cfaa76ecd9, points: 0, grade: false, solution: true -->
```python
def free_prisoners(prison):
  if prison[0] == 0:
    return 0
  count = 1
  for i in range(1, len(prison)):
    if prison[i-1] != prison[i]:
      count += 1
  return count

```

<!-- grade_id: cell-c50554192688359f, points: 2, grade: true, solution: false -->
```python
# Test cases
assert(free_prisoners([1, 1, 1])== 1)
assert(free_prisoners([0, 0, 0])== 0)
assert(free_prisoners([1, 1, 0, 0, 0, 1, 0])== 4)
print('ok')
```

Output:
```
ok
```

**Q5**: Write thr function named "find_possible_pincodes(pin)" that returns a list of all possible variations of a given pincode.

Given a keypad that has the following layout:

<img src="https://drive.google.com/uc?id=1ANXeH2bCl4pWstEJcg3kTIXytvNIh11n" height="30"> 

You will receive a pincode from your secret agent. However, he said it's possible that each of the digits, he was engaged in spying, could be another adjacent digit (horizontally or vertically, but not diagonally).

Assume you have the code, which the secret agent gave you: 46:\
1, 5, or 7 could also be used in place of the 4.\
3, 5, or 9 could also be used in place of the 6.\
Thus, the possible pincodes could be an item in the list of\ ["13","15","16","19","43","45","46","49","53","55","56","59","73","75","76","79"] 

Example: \
-----Argument-----\
"0"\
-----Return--------\
['0', '8']

Example: \
-----Argument-----\
"46"\
-----Return--------\
["13","15","16","19","43","45","46","49","53","55","56","59","73","75","76","79"]

<!-- grade_id: cell-ef0cdf2d0a302112, points: 0, grade: false, solution: true -->
```python
import itertools

def find_adjacent_number(n):
  adjacent_numbers = {
    1: [2,4],
    2: [1,3,5],
    3: [2,6],
    4: [1,5,7],
    5: [2,4,6,8],
    6: [3,5,9],
    7: [4,8],
    8: [5,7,9,0],
    9: [6,8],
    0: [8]
  }  

  return adjacent_numbers[n]

def permutation(n1, n2):
  permutations = list(itertools.product(n1, n2))
  permutations = [int(''.join(map(str, p))) for p in permutations]
  return [str(p) for p in sorted(permutations)]

def possible_pincode(pin):
  if len(pin) == 1:
    return [str(p) for p in sorted(find_adjacent_number(int(pin[0])) + [int(pin)])]
  else:
    digits = [find_adjacent_number(int(c)) + [int(c)] for c in list(str(pin))]

    # merge these two lists
    permuations = permutation(digits[0], digits[1])
    return permuations

```

<!-- grade_id: cell-18522bf2f8cb59c4, points: 2, grade: true, solution: false -->
```python
# Test cases
assert possible_pincode("0") == ['0', '8']
assert possible_pincode("5"), ['2', '4', '5', '6', '8']
assert possible_pincode("46"), ["13","15","16","19","43","45","46","49","53","55","56","59","73","75","76","79"]
print("ok")
```

Output:
```
ok
```